<problem display_name="Custom JavaScript Display and Grading" markdown="null" showanswer="never">
       <customresponse cfn="check_function">
<script type="loncapa/python">

import json
def check_function(e, ans):
  """
  "response" is a dictionary that contains two keys, "answer" and "state".
  The value of "answer" is the JSON string that "getGrade" returns.
  The value of "state" is the JSON string that "getState" returns.
  Clicking either "Submit" or "Save" registers the current state.
  """
  response = json.loads(ans)

  # You can use the value of the answer key to grade:
  answer = json.loads(response["answer"])
  return answer == "correct"

  # Or you can use the value of the state key to grade:
  """
  state = json.loads(response["state"])
  return state["selectedChoice"] == "correct"
  """

</script>
         <p>The input below is an iframe embedding a resource from a third party location that returns a grade back to the platform. You can use this to build custom JavaScript activities and interactive exercises and allow them to be graded.</p>
      <jsinput gradefn="JSInputDemo.getGrade" get_statefn="JSInputDemo.getState" set_statefn="JSInputDemo.setState" initial_state="{&quot;selectedChoice&quot;: &quot;incorrect1&quot;, &quot;availableChoices&quot;: [&quot;incorrect1&quot;, &quot;correct&quot;, &quot;incorrect2&quot;]}" width="600" height="100" html_file="https://files.edx.org/custom-js-example/jsinput_example.html" title="Dropdown with Dynamic Text" sop="false"/>
    </customresponse>
</problem>
