<problem display_name="Molecule Editor" graceperiod="18000 seconds" markdown="null" rerandomize="never" showanswer="attempted" start="2013-02-05T00:00:00Z" weight="null" xqa_key="qaijS3UatK020Wc0sfCtFe0V6jpB4d64">
  <p> La molécule de la Dopamine, comme montrée ci-dessous, ne peut faire de liaison ionique. Éditez la molécule de Dopamine de manière à ce qu'elle puisse faire des liens. 
    </p>
  <p>
        Quand vous êtes prêts, clickez sur "Vérifier.” Si vous voulez recommencer, cliquez sur "Reset.”
    </p>
  <script type="loncapa/python">
def check1(expect, ans):
    import json
    mol_info = json.loads(ans)["info"]
    return any(res == "Can Make Ionic Bonds" for res in mol_info)
    </script>
  <customresponse cfn="check1">
    <editamoleculeinput file="/static/molecules/dopamine.mol">
        </editamoleculeinput>
  </customresponse>
  <solution>
    <img src="/static/Screen_Shot_2013-04-16_at_1.43.36_PM.png"/>
  </solution>
</problem>
