<problem display_name="Molecule Editor" graceperiod="0 day 5 hours 0 minutes 0 seconds" rerandomize="never" showanswer="attempted" start="2013-02-05T00:00" weight="" xqa_key="qaijS3UatK020Wc0sfCtFe0V6jpB4d64">
  <p> The Dopamine molecule, as shown, cannot make ionic bonds. Edit the dopamine molecule so it can make ionic bonds. 
    </p>
  <p>
        When you are ready, click "Check.” If you need to start over, click
        "Reset.”
    </p>
  <script type="loncapa/python">
def check1(expect, ans):
    import json
    mol_info = json.loads(ans)["info"]
    return any(res == "Can Make Ionic Bonds" for res in mol_info)
    </script>
  <customresponse cfn="check1">
    <editamoleculeinput file="/static/molecules/dopamine.mol">
        </editamoleculeinput>
  </customresponse>
  <solution>
    <img src="/c4x/edX/DemoX/asset/Screen_Shot_2013-04-16_at_1.43.36_PM.png"/>
  </solution>
</problem>
